```python
# Initialization cell
try:  # for CS1302 JupyterLite pyodide kernel
    import piplite

    with open("requirements.txt") as f:
        for package in f:
            package = package.strip()
            print("Installing", package)
            await piplite.install(package)
except ModuleNotFoundError:
    pass

import jupytext
import otter

grader = otter.Notebook("main.ipynb")
```

## Lab 0

CS1302 Introduction to Computer Programming
___

## Setup

### JupyterHub

#### How to access the JupyterHub server?

- Enter the url of the Jupyterhub server [divedeep.cs.cityu.edu.hk/cs1302](https://divedeep.cs.cityu.edu.hk/cs1302) in a web browser.
- Enter your [EID](https://www.cityu.edu.hk/esu/eid.htm) and Password in the fields `Username` and `Password` respectively.
- Click the `Sign in` button.  

  ![log in](images/login.dio.svg)

- If you are presented with the choices of servers, choose `CS1302` as the server option and click `Start`.  

  ![server options](images/serveroptions.dio.svg)

```{note}

`Minimal` is only for trouble-shooting and is insufficient for running the course materials. If you accidentally choose it, you can restart your server after entering the JupyterLab interface as described below.

```

- The JupyterLab interface should appear.  

  ![jupyterlab](images/jupyterlab.dio.svg)

```{tip}

**How to troubleshoot in case or error?**

- Refresh the page, re-login, or clear your browser cache to reset your browser.
- Restart your server using the `Hub Control Panel` under the `File` menu (shown above) or visit the hub home  
  https://divedeep.cs.cityu.edu.hk/cs1302/hub/home  
  and click `Stop My Server` and then `Start My Server`.  
  ![Restart my server](images/restart.dio.svg)
- Contact your instructor using the contact information on the [course homepage](https://canvas.cityu.edu.hk/courses/42868).

```

### Jupyter Notebook

#### How to access course materials?

The recommended way is to follow the links to the lecture/lab notebooks on the [course homepage](https://canvas.cityu.edu.hk/courses/48943).

![course homepage](images/canvas.dio.svg) 

```{important}

The recommended way to access the notebook is to use [`nbgitpuller`](https://jupyterhub.github.io/nbgitpuller/index.html) links which

- [git-pull](https://git-scm.com/docs/git-pull) updates/corrections to the notebooks and 
- merge them with the notebooks stored under `cs1302` folder in your home directory but
- without overwriting your changes.

If you open the notebook directly on JupyterHub instead of following the git-pull link, updates/corrections to notebooks, if any, will not be pulled.

```

#### How to complete a lab assignment?

Learn how to edit a notebook:

1. Click `Help`$\to$ `Welcome Tour` to learn the JupyterLab interface.  
2. Click `Help`$\to$ `Notebook Tour` to learn about the jupyter notebook.
3. Watch the official [video tutorial](https://youtu.be/A5YyoCKxEOU) to JupyterLab.

**Exercise** 

In learning a new programming language, the first program to write is often the ["Hello, World!"](https://en.wikipedia.org/wiki/%22Hello,_World!%22_program) program, which says Hello to the world.  

Type the program below and run it with `Shift+Enter`. Use <kbd>tab</kbd> or 4 spaces to indent.

![Hello world program](images/hello.dio.svg) 

<!-- grade_id: Hello-World, points: 0, grade: false, solution: true -->
```python
def say_hello():
    ...
    
say_hello()
```

In order to check your work thoroughly, there will be both *visible* and *hidden* test cases. The following is a visible test you can run to check your answer:

```python
grader.check("Hello-World")
```

The test returns an assertion error if your program does not print the correct message.

```{tip}

1. You can repeatedly modify your solution and run the test cell until your solution passes the test. There is no mark penalty in failing the test before submission. 
1. You should try to understand the error messages from a failed test.
1. To assess your solution thoroughly, we will run new tests hidden from you after you have submitted your notebook. Therefore, *you should ensure your solution works in general rather than just the visible tests*.
1. If you open the same notebook multiple times in different browser windows, be careful in making changes in different windows as you may overwrite your own changes.
1. If your notebook fails to run any code, the kernel might have died. You can restart the kernel with `Kernel`$\to$ `Restart`. If restarting fails, check your code cells to see if there is anything that breaks the kernel.

```

#### How to submit a notebook

Lecture notebooks under the subfolders `Lecture*/` need NOT be submitted. You only need to submit the lab notebooks under `Lab*\`. Although the first Lab (`Lab0`) does not count towards your final grade, you should still submit it, to get familiar with the procedure.

```{important}

Before you submit, make sure everything runs as expected:

1. **Git-pull the notebooks** (optional): Follow any one of the link on the course homepage to a notebook, which will git-pull any updates/corrections to (all) your notebooks.
1. **Save the notebook**: Unsaved changes are not exported, even though they are in the memory.
1. **Restart the kernel**: `Kernel`$\to$ `Restart Kernel...` to have a clean state before running visible tests.
1. **run all cells**: `Run`$\to$ `Run All Cells` which  
    - not only checks you code, but also
    - also generates 
        - `main.py` file for graders to check your code, and
        - a zip file containing all your files for submission.
1. **Upload the zip file** to the corresponding Canvas assignment page.
1. **Enter a pseudonym**, if required, for the leaderboard.
```

```{attention}

Grading policies:
- *No late submission* will be collected without valid justifications.
- *Double check* that you have submitted the correct Lab assignment.
- You can repeatedly submit before the deadline without penalty, so submit your assignment in advance. You are responsible for 
  - failed submission due to network outage too close to the deadline; and
  - make a record or backup of your submission with the submission timestamp.
  
You submission may not be graded in the following cases:

- The notebook files are renamed, e.g., `main.ipynb` changed/copied to `main (Copy).ipynb`.
- `main.py` does not contain your code in `main.ipynb`. `main.py` should be automatically generated for similarity check against code plagiarism.
- An html file exists with the same name as a notebook, e.g., `Setup.html`.
- The file size is too large, e.g, over `100MB`.
- The code takes too long to run, e.g., over `10mins` or requires an excessive amount memory, e.g., over `500MB`.
```

```{tip}

If you believe your notebooks are corrupted, 
1. download/backup your existing notebooks,
1. remove them from the `Lab*/` folder,
1. click the git-pull links from the course homepage to re-pull the notebooks, and
1. copy your solutions to the new notebooks.

```

```{seealso}

You may also run the course notebooks outside the jupyterhub server and locally on your computer. For more details, see the course homepage.

```

## Debugger

### Why debuggers?

Computer programs can somtimes behave in an unexpected way due to *Errors* or *Bugs*. Finding the errors might be easy for simple programs. However, for programs consisting of many lines of code, finding errors can be very hard.

A *Debugger* allows us to pause/break the execution of a program at a line, called a *breakpoint*, we specify. This allows us to find errors more easily.

 In this lab, we will learn to use a visual debugger in JupyterLab. 

### Set breakpoints

To use the Jupyterlab visual debugger, ensure XPython kernel is running:

![xpython](images/xpython.dio.svg) 

If you do not see XPython or a circle next to it on the top-right-hand corner, switch to the XPython kernel by clicking the name of the current kernel on the top-right-hand corner.

We will use the following code as an example to set a breakpoint. First, execute the code to see the effect.

```python
msg = "Hello, World!"
print(msg)
```

To enter the debug mode, turn on the switch near on top right-hand corner: 

![jupyterlab debugger switch](images/jupyterlab-debug.dio.svg)

The debugging sidebar should slide open. You can also open/close the sidebar by clicking the bug icon on the right. 

**Exercise** 

Code cells also looks different in debug mode. What is the difference?

![code cell](images/codecell.dio.svg)

_Type your answer here, replacing this text._

To set a breakpoint, press the dot that appears as you hover to the left of line number 2. Run the cell to see what happens.

```python
msg = "Hello, World!"
print(msg)
```

The debugger provides quite a bit of information, organized as follows:

* **Orange line** is the line where the execution is paused.  
  ![paused](images/paused.dio.svg)

* **Variables panel** shows the current values of different variables.  
  ![variables](images/variables.dio.svg) 

* **Callstack panel** provides a more detailed context of where the execution is paused. It also provides the flow control buttons to resume/stop the execution.  
  ![callstack](images/callstack.dio.svg)

* **Breakpoints panel** shows all the breakpoints of the current debug session.  
  ![breakpoints](images/breakpoints.dio.svg)

* **Source panel** highlights the source of the breakpoint at which the execution is currently paused.  
  ![source](images/source.dio.svg)

**Exercise** Explain whether line 1 and 2 are executed. Use the information from the debugger to support your answer.

_Type your answer here, replacing this text._

### Resume execution

To resume/terminate the execution, use the following flow control buttons at the top of the callstack panel:  

  ![callstack](images/flowcontrol.dio.svg)

- **Continue** button continues the execution and pause again if it hits a breakpoint.
- **Terminate** button stops the execution.
- **Next/Step-over** button continues the execution of program to the next statement without entering function.

### Run temporary code

Sometimes, we want to run temporary code without modifying the current notebook. To do so:

- Right click anywhere on a notebook and choose `New Console for Notebook`:

  ![new console](images/newconsole.dio.svg) 

- Type some code such as `msg` into the console input and run it with `Shift-Enter`:

  ![console](images/console.dio.svg) 

**Exercise** 

What does the following function `f(n)` computes? Try to set a breakpoint and use the flow control to understand how the program works.

```python
def f(n): # definition of a function
    return n * f(n - 1) if n > 0 else 1


f(3)
```

```{hint}

Set the breakpoint one at a time but at different locations. If you set your breakpoint at line 5, you should use the **Step In** button to step into the execution of the function.

You can also make use of the new console for notebook:
- You can run `f(n)` for different values of `n` in the console to understand how the function works. 
- You can set a breakpoint in Line 2 of the previous code cell containing `def f(n): ...` and run `f(3)` in the console to pause at the breakpoint. The source code containing the breakpoint will also be shown in the **Source panel**. 
- You can also set a breakpoint in the console input, say `f(3)`, and then step into it.
- While the execution is paused, you can enter a new line of code such as `print(msg)` below `f(3)` in the console and then continue the execution to run the new line of code.

```

_Type your answer here, replacing this text._

```{seealso}

There are other debuggers in different interfaces. You may use the vscode interface in our JupyterHub server:  
- `File`$\to$ `New Launcher`$\to$ `VS Code`

Follow the [official guide](https://code.visualstudio.com/docs/datascience/jupyter-notebooks) to install the python/jupyter extension and debug a jupyter notebook. 

```

### Visualize code execution

Seeing is believing. You can also debug your code by visualizing the code execution.

First, load the `divewidgets` using a line magic:

```python
%reload_ext divewidgets
```

Then, create a cell magic with the python code you want to visualize:

```python
%%optlite --height 500
def f(n): # definition of a function
    return n * f(n - 1) if n > 0 else 1


f(3)
```

Try the following:
- Drag the slider to time-travel the code execution.
- Click any lines with code to toggle breakpoints, and click `Next >` (`< Prev`) to go the the next (previous) breakpoint.
- Click the `Edit the code` link to edit the code.
- Click the `Permalink` button to generate a url of the code, and copy and paste the url to a new browser window to run the visualization there.
- Click the `Live Edit` button to live edit the code in a new browser window, where your changes to the code are immediately visualized. You can also time-travel and generate a permalink there.

Run the following cell to see additional options available to the cell magic.

```python
%%optlite?
```

```{tip}

We have developed [OPTLite](https://github.com/dive4dec/optlite) to make the popular learning tool [Online Python Tutor](https://pythontutor.com/) run locally on your browser without a server. The serverless setup work offline once the page is loaded with no advertisements/restrictions on your code. 

If you have difficulty understanding or debugging a piece of code yourself, use OPTLite to share a link of the code with others. An example of a permalink is:

> https://dive4dec.github.io/optlite/#code=def%20f%28n%29%3A%20%23%20definition%20of%20a%20function%0A%20%20%20%20return%20n%20*%20f%28n%20-%201%29%20if%20n%20%3E%200%20else%201%0A%0A%0Af%283%29&mode=edit&origin=opt-frontend.js&rawInputLstJSON=%5B%5D

If `mode=edit` in the query string is changed to `mode=display`, the visualization will start automatically:

> https://dive4dec.github.io/optlite/#code=def%20f%28n%29%3A%20%23%20definition%20of%20a%20function%0A%20%20%20%20return%20n%20*%20f%28n%20-%201%29%20if%20n%20%3E%200%20else%201%0A%0A%0Af%283%29&mode=display&origin=opt-frontend.js&rawInputLstJSON=%5B%5D

To enter the live editing mode automaticaslly, specify `live.html` to the path:

> https://dive4dec.github.io/optlite/live.html#code=def%20f%28n%29%3A%20%23%20definition%20of%20a%20function%0A%20%20%20%20return%20n%20*%20f%28n%20-%201%29%20if%20n%20%3E%200%20else%201%0A%0A%0Af%283%29&mode=edit&origin=opt-frontend.js&rawInputLstJSON=%5B%5D

```

**Exercise**

Give a permalink that visualize the Hello World program immediately without the need to click `Visualize` button. 

_Type your answer here, replacing this text._

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# extra files to submit
extra_files = []
```

```python
# Generate the source main.py necessary for grading and similarity check.
jupytext.write(jupytext.read("main.ipynb"), "main.py", fmt="py:percent")
```

```python
# Generate the zip file to submit.
grader.export(pdf=False, run_tests=False, files=["main.py", *extra_files])
```

Tests for Hello-World (points: None)
```
>>> # Run this test cell right after running your "Hello, World!" program.
>>> import io
>>> import sys
>>> 
>>> old_stdout, sys.stdout = sys.stdout, io.StringIO()
>>> say_hello()
>>> printed = sys.stdout.getvalue()
>>> sys.stdout = old_stdout
>>> assert printed == "Hello, World!\n"
```
