```python
# Initialization cell
try:  # for CS1302 JupyterLite pyodide kernel
    import piplite

    with open("requirements.txt") as f:
        for package in f:
            package = package.strip()
            print("Installing", package)
            await piplite.install(package)
except ModuleNotFoundError:
    pass

import math
import jupytext
import numpy as np
import otter
from ipywidgets import interact

grader = otter.Notebook("main.ipynb")
```

# Lab 3a

**CS1302 Introduction to Computer Programming**
___

In this notebook, we will improve the quadratic equation solver in the previous lab using conditional executions.

## Discriminant

Recall that a quadratic equation is

$$
ax^2+bx+c=0
$$

where $a$, $b$, and $c$ are real-valued coefficients, and $x$ is the unknown variable.

---

**Definition** (Discriminant)

The disciminant of a quadratic equation is defined as

$$
\Delta := b^2 - 4ac.
$$ (discriminant)

---

The discriminant $\Delta$ discriminates the roots 

$$\frac{-b\pm \sqrt{\Delta}}{2a}$$ (roots)

of a quadratic equation:

---

**Proposition**

The roots of a quadratic equation {eq}`quadratic` are both equal to 

$$-\frac{b}{2a}$$ (repeated root)

if and only if the discriminant is zero, i.e., $\Delta=0$.

---

For example, if $(a,b,c)=(1,-2,1)$, then 

$$\Delta = (-2)^2 - 4(1\cdot 1) = 0$$

and so the repeated root is

$$
- \frac{b}{2a} = \frac{2}2 = 1.
$$

$(x-1)^2 = x^2 - 2x + 1$ as expected.

**Exercise**  (zero-determinant)

Complete the following function by assigning `roots` only one root when the discriminant is zero. E.g., if `(a, b, c) == (1, -2, 1)`, then `roots` should be assigned the value `1.0` instead of `1.0, 1.0`.

```{note}
Replace `...` in the following code cell by your code. If you have a better implementation, you need *not* follow the solution template. E.g., you can write your solution entirely using boolean operations without any `if` statement.
```


```python
import math
```

<!-- grade_id: zero_determininant, points: 0, grade: false, solution: true -->
```python
def get_roots(a, b, c):
    d = b ** 2 - 4 * a * c  # discriminant
    if math.isclose(d, 0):
        roots = ...  # repeated root
    else:
        d **= 0.5
        roots = ..., ...  # distinct roots
    return roots
```

```python
grader.check("zero-determinant")
```

**Exercise** (isclose)

Why use `math.isclose(d, 0)` instead of `d == 0`?

_Type your answer here, replacing this text._

## Linear equation

**Exercise** (zero-a)

Give the name of the error that the formula {eq}`roots` implemented in python would when $a=0$. Explain whether it is a runtime error or not.

_Type your answer here, replacing this text._

Nevertheless, the quadratic equation remains valid:

$$
bx + c=0,
$$

the root of which is $x = -\frac{c}b$.

**Exercise** (linear)

Improve the function `get_roots` to return the root $-\frac{c}{b}$ when $a=0$.

```python
import math
```

<!-- grade_id: linear, points: 0, grade: false, solution: true -->
```python
def get_roots(a, b, c):
    d = b**2 - 4 * a * c  # discriminant
    if ...:  # linear equation
        roots = ...
    elif math.isclose(d, 0):
        roots = ...  # repeated root
    else:
        d **= 0.5
        roots = ...  # distinct roots
    return roots
```

```python
grader.check("linear")
```

## Degenerate cases

What if $a=b=0$? In this case, the equation becomes

$$
c = 0
$$
which is always satisfied if $c=0$ but never satisfied if $c\neq 0$. 

**Exercise** (degenerate)

Improve the function `get_roots` to return root(s) under all the following cases:
- If $a=0$ and $b\neq 0$, assign `roots` to the single root $-\frac{c}{b}$. 
- If $a=b=0$ and $c\neq 0$, assign `roots` to `None`.
- If $a=b=c=0$, there are infinitely many roots. Assign to `roots` the tuple `-float('inf'), float('inf')`.  
    Note that `float('inf')` converts the string `'inf'` to a floating point value that represents $\infty$.
    
```{caution}
- `None` a *python keyword* that refers to a special object. You should *NOT* regard it as a string or quote it. (~~`"None"`~~).
- `float('inf')` converts the string `'inf'` to a floating point number no smaller than any other floating point numbers. You should *NOT* write `inf`, which needs not refer to `float('int')`.
```

```python
import math
```

<!-- grade_id: degenerate, points: 0, grade: false, solution: true -->
```python
def get_roots(a, b, c):
    d = b ** 2 - 4 * a * c
    if ...:  # linear
        if ...:  # degenerate
            if ...:
                roots = ..., ...
            else:  # no roots
                roots = ...
        else:
            roots = ...
    elif math.isclose(d, 0):
        roots = ...  # repeated root
    else:
        d **= 0.5
        roots = ...  # distinct roots
    return roots
```

```python
grader.check("degenerate")
```

After you have complete the exercises, you can run your robust solver below:

```python
# quadratic equations solver
@interact(a=(-10, 10, 1), b=(-10, 10, 1), c=(-10, 10, 1))
def quadratic_equation_solver(a=1, b=2, c=1):
    print("Root(s):", get_roots(a, b, c))
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# extra files to submit
extra_files = []
```

```python
# Generate the source main.py necessary for grading and similarity check.
jupytext.write(jupytext.read("main.ipynb"), "main.py", fmt="py:percent")
```

```python
# Generate the zip file to submit.
grader.export(pdf=False, run_tests=False, files=["main.py", *extra_files])
```

Tests for degenerate (points: None)
```
>>> import numpy as np
>>> 
>>> assert np.isclose(get_roots(1, 1, 0), (-1.0, 0.0)).all()
>>> assert np.isclose(get_roots(1, 0, 0), 0.0).all()
>>> assert np.isclose(get_roots(1, -2, 1), 1.0).all()
>>> assert np.isclose(get_roots(0, -2, 1), 0.5).all()
>>> assert np.isclose(get_roots(0, 0, 0), (-float('inf'), float('inf'))).all()
>>> assert get_roots(0, 0, 1) is None
```


Tests for linear (points: None)
```
>>> import numpy as np
>>> 
>>> assert np.isclose(get_roots(1, 1, 0), (-1.0, 0.0)).all()
>>> assert np.isclose(get_roots(1, 0, 0), 0.0).all()
>>> assert np.isclose(get_roots(1, -2, 1), 1.0).all()
>>> assert np.isclose(get_roots(0, -2, 1), 0.5).all()
```


Tests for zero-determinant (points: None)
```
>>> import numpy as np
>>> 
>>> assert np.isclose(get_roots(1, 1, 0), (-1.0, 0.0)).all()
>>> assert np.isclose(get_roots(1, 0, 0), 0.0).all()
```
