```python
# Initialization cell
try:  # for CS1302 JupyterLite pyodide kernel
    import piplite

    with open("requirements.txt") as f:
        for package in f:
            package = package.strip()
            print("Installing", package)
            await piplite.install(package)
except ModuleNotFoundError:
    pass

import random
import jupytext
import otter
from ipywidgets import interact

grader = otter.Notebook("main.ipynb")
%reload_ext divewidgets
```

# Pandas

**CS1302 Introduction to Computer Programming**
___

In this lab, we will analyze COVID19 data using a powerful package called [`pandas`](https://pandas.pydata.org/docs/user_guide/index.html).  
The package name comes from *panel data* and *Python for data analysis*.

## Loading CSV Files with Pandas 

[DATA.GOV.HK](https://data.gov.hk/en-data/dataset/hk-dh-chpsebcddr-novel-infectious-agent) provides an [API](https://data.gov.hk/en/help/api-spec#historicalAPI) to retrieve historical data on COVID-19 cases in Hong Kong.

The following uses the `urlencode` function to create the url that links to a csv file containing probable and confirmed cases of COVID-19 by Aug 1st, 2020.

```python
from urllib.parse import urlencode

url_data_gov_hk_get = "https://api.data.gov.hk/v1/historical-archive/get-file"
url_covid_csv = "http://www.chp.gov.hk/files/misc/enhanced_sur_covid_19_eng.csv"
time = "20200801-1204"
url_covid = url_data_gov_hk_get + "?" + urlencode({"url": url_covid_csv, "time": time})

print(url_covid)
```

Output:
```
https://api.data.gov.hk/v1/historical-archive/get-file?url=http%3A%2F%2Fwww.chp.gov.hk%2Ffiles%2Fmisc%2Fenhanced_sur_covid_19_eng.csv&time=20200801-1204
```

```{note}

`urlencode` creates a string `'url=<...>&time=<...>'` with some [special symbols encoded](https://www.w3schools.com/tags/ref_urlencode.ASP), e.g.:
- `:` is replaced by `%3A`, and
- `/` is replaced by `%2F`.

```

**Exercise** (simple-encode)

Write a function `simple_encode` that takes in a string and return a string with `:` and `/` encoded as described above.

```{hint}

Use the `replace` method of `str`.

```

<!-- grade_id: simple_encode, points: 0, grade: false, solution: true -->
```python
def simple_encode(string):
    return string.replace(":", "%3A").replace("/", "%2F")
```

```python
grader.check("simple-encode")
```

Output:
```
simple-encode results: All test cases passed!
```

Like the function `open` that loads a file into memory, `pandas` has a function `read_csv` that loads a csv file. The csv file can even reside on the web:

```python
import pandas as pd

df_covid = pd.read_csv(url_covid)

print(type(df_covid))
df_covid
```

Output:
```
<class 'pandas.core.frame.DataFrame'>
      Case no. Report date Date of onset Gender  Age  \
0            1  23/01/2020    21/01/2020      M   39   
1            2  23/01/2020    18/01/2020      M   56   
2            3  24/01/2020    20/01/2020      F   62   
3            4  24/01/2020    23/01/2020      F   62   
4            5  24/01/2020    23/01/2020      M   63   
...        ...         ...           ...    ...  ...   
3268      3269  31/07/2020    26/07/2020      M   22   
3269      3270  31/07/2020    28/07/2020      F   31   
3270      3271  31/07/2020  Asymptomatic      F   36   
3271      3272  31/07/2020       Pending      F   22   
3272      3273  31/07/2020    28/07/2020      M   68   

      Name of hospital admitted Hospitalised/Discharged/Deceased  \
0                           NaN                       Discharged   
1                           NaN                       Discharged   
2                           NaN                       Discharged   
3                
```

```{note}

The above creates a [`DataFrame` object](https://pandas.pydata.org/docs/reference/api/pandas.DataFrame.html?highlight=dataframe#pandas.DataFrame):  
- The content of the csv file is displayed as an HTML table conveniently.   
- We can control how much information to show by setting the [display options](https://pandas.pydata.org/pandas-docs/stable/user_guide/options.html).)

```

**Exercise** (building)

Using the function `pd.read_csv`, load `building_list_eng.csv` as `df_building` from the url `url_building`.  

<!-- grade_id: df_building, points: 0, grade: false, solution: true -->
```python
url_building_csv = "http://www.chp.gov.hk/files/misc/building_list_eng.csv"
time = "20200801-1203"
url_building = (
    url_data_gov_hk_get + "?" + urlencode({"url": url_building_csv, "time": time})
)
df_building = pd.read_csv(url_building)
df_building
```

Output:
```
District                                      Building name  \
0         Yau Tsim Mong                                      David Mansion   
1         Yau Tsim Mong                                  Kimberley Mansion   
2          Wong Tai Sin                                        Gomme House   
3                 North                   Yung Sui House, Yung Shing Court   
4          Sham Shui Po                                    Shun King Court   
...                 ...                                                ...   
2160            Sha Tin                MOSTown (Phase 4) (non-residential)   
2161           Tuen Mun                   San Hui Market (non-residential)   
2162       Sham Shui Po  Shek Kip Mei Estate Commercial Centre (non-res...   
2163  Central & Western       Federate Building Fairwood (non-residential)   
2164  Central & Western                    Pacific Place (non-residential)   

     Last date of residence of the case(s)  \
0                               
```

```python
grader.check("building")
```

Output:
```
building results: All test cases passed!
```

## Selecting and Removing columns

We can obtain the column labels of a `Dataframe` using its `columns` attribute.

```python
df_covid.columns
```

Output:
```
Index(['Case no.', 'Report date', 'Date of onset', 'Gender', 'Age',
       'Name of hospital admitted', 'Hospitalised/Discharged/Deceased',
       'HK/Non-HK resident', 'Case classification*', 'Confirmed/probable'],
      dtype='object')
```

Using the indexing operator `[]`, a column of a `DataFrame` can be returned as a [`Series` object](https://pandas.pydata.org/pandas-docs/stable/reference/api/pandas.Series.html), which is essentially a named array.   
We can further use the method `value_counts` to return the counts of different values in another `Series` object.

```python
series_gender_counts = df_covid[
    "Gender"
].value_counts()  # return the number of male and female cases

print(type(series_gender_counts))
series_gender_counts
```

Output:
```
<class 'pandas.core.series.Series'>
F    1648
M    1625
Name: Gender, dtype: int64
```

**Exercise** (district)

For `df_building`, use the operator `[]` and method `value_counts` to assign `series_district_counts` to a `Series` object that stores the counts of buildings in different district.

<!-- grade_id: series_district_counts, points: 0, grade: false, solution: true -->
```python
series_district_counts = df_building["District"].value_counts()
```

```python
grader.check("district")
```

Output:
```
district results: All test cases passed!
```

In `df_covid`, it appears that the column `Name of hospital admitted` contains no information. We can confirm this by
1. returning the column as a `Series` with `df_covid_cases['Name of hospital admitted']`, and
1. printing an array of unique column values using the method `unique`.

```python
df_covid["Name of hospital admitted"].unique()
```

Output:
```
array([nan])
```

**Exercise** (drop)

Drop the column `Name of hospital admitted` from `df_covid` using the `drop` method of the DataFrame. 

```{hint}

Consider reading the documentation of the `drop` method for  
- mutating `df_covid` in place instead of creating a copy of the DataFrame with the column dropped, but
- suppressing error when dropping a column that does not exist or has already been dropped.

```

<!-- grade_id: drop, points: 0, grade: false, solution: true -->
```python
df_covid.drop(columns="Name of hospital admitted", inplace=True)
```

```python
grader.check("drop")
```

Output:
```
drop results: All test cases passed!
```

## Selecting Rows of DataFrame

We can select the confirmed male cases using the attribute [`loc`](https://pandas.pydata.org/pandas-docs/stable/reference/api/pandas.DataFrame.loc.html) and the indexing operator `[]`.

```python
df_confirmed_male = df_covid.loc[
    (df_covid["Confirmed/probable"] == "Confirmed") & (df_covid["Gender"] == "M")
]
print(type(df_covid.loc))
df_confirmed_male
```

Output:
```
<class 'pandas.core.indexing._LocIndexer'>
      Case no. Report date Date of onset Gender  Age  \
0            1  23/01/2020    21/01/2020      M   39   
1            2  23/01/2020    18/01/2020      M   56   
4            5  24/01/2020    23/01/2020      M   63   
5            6  26/01/2020    21/01/2020      M   47   
7            8  26/01/2020    25/01/2020      M   64   
...        ...         ...           ...    ...  ...   
3264      3265  31/07/2020    25/07/2020      M   54   
3265      3266  31/07/2020    30/07/2020      M   69   
3267      3268  31/07/2020    23/07/2020      M   61   
3268      3269  31/07/2020    26/07/2020      M   22   
3272      3273  31/07/2020    28/07/2020      M   68   

     Hospitalised/Discharged/Deceased HK/Non-HK resident  \
0                          Discharged    Non-HK resident   
1                          Discharged        HK resident   
4                          Discharged    Non-HK resident   
5                          Discharged       
```

```{note}

`loc` essentially returns an object that implements the [advanced indexing method](https://numpy.org/doc/stable/user/basics.indexing.html#advanced-indexing) for `__getitem__`. In particular, the above uses [boolean indexing](https://numpy.org/doc/stable/user/basics.indexing.html#boolean-array-indexing).
```

**Exercise** (local)

Assign `df_confirmed_local` to a `DataFrame` of confirmed cases that are local or epidemiologically linked with a local case.

<!-- grade_id: df_confirmed_local, points: 0, grade: false, solution: true -->
```python
df_confirmed_local = df_covid.loc[
    (df_covid["Confirmed/probable"] == "Confirmed") &
    ((df_covid["Case classification*"] == "Local case") |
        (df_covid["Case classification*"] == "Epidemiologically linked with local case"))
]
df_confirmed_local
```

Output:
```
Case no. Report date Date of onset Gender  Age  \
15          16  04/02/2020    23/01/2020      F   64   
16          17  04/02/2020    22/01/2020      M   60   
17          18  04/02/2020    28/01/2020      M   25   
18          19  05/02/2020    30/01/2020      F   28   
19          20  05/02/2020    04/02/2020      F   56   
...        ...         ...           ...    ...  ...   
3268      3269  31/07/2020    26/07/2020      M   22   
3269      3270  31/07/2020    28/07/2020      F   31   
3270      3271  31/07/2020  Asymptomatic      F   36   
3271      3272  31/07/2020       Pending      F   22   
3272      3273  31/07/2020    28/07/2020      M   68   

     Hospitalised/Discharged/Deceased HK/Non-HK resident  \
15                         Discharged        HK resident   
16                         Discharged        HK resident   
17                         Discharged        HK resident   
18                         Discharged        HK resident   
19                         Discha
```

```python
grader.check("local")
```

Output:
```
local results: All test cases passed!
```

**Exercise** (case-counts)

Write a function `case_counts` that 
- takes an argument `district`, and
- returns the number of cases in `district`. 

```{hint}

Be careful that there can be more than one case for each building and there may be multiple buildings associated with one case. You may want to use the `split` and `strip` methods of `str` to obtain a list of cases from the `Dataframe`.

```

<!-- grade_id: case_counts, points: 0, grade: false, solution: true -->
```python
def case_counts(district):
    # print(df_building["District"].value_counts())
    df_district_list = df_building.loc[
        df_building["District"] == district
    ]
    # print(df_district_list)
    # df_cases_set = df_district_list["Related probable/confirmed cases"]
    cases_set = set()
    for cases in df_district_list["Related probable/confirmed cases"]:
        # print(cases.split(","))
        for case in cases.split(","):
            cases_set.add(case.strip())
    return len(cases_set)
# case_counts("Central & Western")
```

```python
grader.check("case-counts")
```

Output:
```
case-counts results: All test cases passed!
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# extra files to submit
extra_files = []
```

```python
# Generate the source main.py necessary for grading and similarity check.
jupytext.write(jupytext.read("main.ipynb"), "main.py", fmt="py:percent")
```

```python
# Generate the zip file to submit.
grader.export(pdf=False, run_tests=False, files=["main.py", *extra_files])
```

Output:
```
<IPython.core.display.HTML object>
```

```python

```

Tests for building (points: None)
```
>>> # tests
>>> assert all(df_building.columns == [
...     "District",
...     "Building name",
...     "Last date of residence of the case(s)",
...     "Related probable/confirmed cases"])  # check column names
```


Tests for case-counts (points: None)
```
>>> # tests
>>> assert case_counts("Kwai Tsing") == 109
```


Tests for district (points: None)
```
>>> # tests
>>> assert all(series_district_counts[["Wong Tai Sin", "Kwun Tong"]] == [313, 212])
```


Tests for drop (points: None)
```
>>> # tests
>>> assert all(df_covid.columns == [
...         "Case no.",
...         "Report date",
...         "Date of onset",
...         "Gender",
...         "Age",
...         "Hospitalised/Discharged/Deceased",
...         "HK/Non-HK resident",
...         "Case classification*",
...         "Confirmed/probable"])
```


Tests for local (points: None)
```
>>> # tests
>>> assert set(df_confirmed_local["Case classification*"].unique()) == {
...     "Epidemiologically linked with local case",
...     "Local case"}
```


Tests for simple-encode (points: None)
```
>>> # tests
>>> assert simple_encode("http://www.chp.gov.hk/files/misc/enhanced_sur_covid_19_eng.csv") == "http%3A%2F%2Fwww.chp.gov.hk%2Ffiles%2Fmisc%2Fenhanced_sur_covid_19_eng.csv"
```
